---
title: "Univariate Description: Ordinal Variables"
tutorial:
  id: "05-Ordinal-Description"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Learn how to derscribe an ordinal variable: central tendency, dispersion, bar plot.
---

```{=html}
<style> /* Only show focus indicators when actually focused */
button:focus-visible, .btn:focus-visible {
  outline: 2px solid #6c757d;
  outline-offset: 1px;
}
/* Remove default focus styles that might conflict but preserve cursor */
button:focus, .btn:focus {
  outline: none;
}
/* Ensure proper cursor states for buttons */
button, .btn {
  cursor: pointer;
}
button:disabled, .btn:disabled {
  cursor: not-allowed;
}
/* Ensure keyboard users can see focus on code editors */
.ace_editor.ace_focus {
  border: 2px solid #6c757d !important;
}
/* Better contrast for exercise areas */
.tutorial-exercise {
  border: 1px solid #f0f0f0;
  border-radius: 4px;
  padding: 15px;
  margin: 20px 0;
}
/* Improved font and readability */
body {
  font-size: 16px;
  line-height: 1.5;
}
/* Accessible replacement for blue text */
.important-note {
  background-color: #e3f2fd;
  border-left: 4px solid #1976d2;
  padding: 12px;
  margin: 15px 0;
  font-weight: bold;
  border-radius: 4px;
}
/* Success/completion messages */
.success-note {
  background-color: #e8f5e8;
  border-left: 4px solid #4caf50;
  padding: 12px;
  margin: 15px 0;
  font-weight: bold;
  border-radius: 4px;
}
/* Package name styling */
.package-name {
  background-color: #e8f5e8;
  color: #2e7d32;
  padding: 2px 4px;
  border-radius: 3px;
  font-weight: bold;
}
.important-text { background-color: #e8f0ff; /* soft blue background (optional) */ color: #1565c0; /* blue text */ padding: 2px 4px; border-radius: 3px; font-weight: bold; }
/* Fix for Run Code button hover states - more specific targeting */
.tutorial-exercise-actions .btn {
  cursor: pointer !important;
}
.tutorial-exercise-actions .btn:hover {
  cursor: pointer !important;
}
.tutorial-exercise-actions .btn:disabled {
  cursor: not-allowed !important;
}
/* Override any conflicting learnr cursor rules */ [data-label] .btn,
.tutorial-exercise .btn,
button[data-label] {
  cursor: pointer !important;
}
/* Accessible table styling */
.table-container {
  overflow-x: auto;
  margin: 20px 0;
}
table {
  width: 100%;
  border-collapse: collapse;
  border: 1px solid #ddd;
}
th, td {
  border: 1px solid #ddd;
  padding: 12px;
  text-align: left;
}
th {
  background-color: #f5f5f5;
  font-weight: bold;
}
caption {
  padding: 10px;
  text-align: left;
  font-size: 1.1em;
  margin-bottom: 10px;
}
/* Focus styles for table cells */
th:focus, td:focus {
  outline: 2px solid #6c757d;
  outline-offset: 1px;
}
/* Skip navigation for keyboard users */
.skip-nav {
  position: absolute;
  top: -40px;
  left: 6px;
  background: #000;
  color: #fff;
  padding: 8px 16px;
  text-decoration: none;
  z-index: 1000;
  border-radius: 4px;
  font-weight: bold;
}
.skip-nav:focus {
  top: 6px;
  outline: 2px solid #fff;
}
/* Enhanced focus visibility for tutorial navigation */
.tutorial-panel-name:focus,
.tutorial-panel-name:focus-visible {
  outline: 1px solid #6c757d !important;
  outline-offset: 1px !important;
  background-color: rgba(108, 117, 125, 0.1) !important;
}
/* Better focus for exercise containers */
.tutorial-exercise:focus-within {
  outline: 1px solid #6c757d;
  outline-offset: 1px;
  background-color: rgba(108, 117, 125, 0.05);
}
/*New*/ /* Enhanced error message accessibility */
.alert-danger, .error, .tutorial-exercise-error {
  border-left: 4px solid #d32f2f;
  padding: 12px;
  margin: 10px 0;
  background-color: #ffebee;
  border-radius: 4px;
}
.alert-success, .tutorial-exercise-success {
  border-left: 4px solid #4caf50;
  padding: 12px;
  margin: 10px 0;
  background-color: #e8f5e8;
  border-radius: 4px;
}
/* Ensure error messages are visible and well-contrasted */
.tutorial-exercise-error {
  color: #d32f2f;
  font-weight: bold;
}
.tutorial-exercise-success {
  color: #2e7d32;
  font-weight: bold;
}
/* Focus styles for elements with error associations */
.ace_editor[aria-describedby]:focus {
  border-color: #d32f2f !important;
  box-shadow: 0 0 0 2px rgba(211, 47, 47, 0.2) !important;
}
</style>
```

```{=html}
<script>
document.addEventListener('DOMContentLoaded', function() {
  
  // Simple, reliable keyboard shortcut help
  document.addEventListener('keydown', function(e) {
    if (e.ctrlKey && e.shiftKey && e.key === '?') {
      e.preventDefault();
      alert(`R Tutorial Navigation Help:

KEYBOARD LIMITATIONS:
- Tab only moves between headings (not buttons or code boxes)
- Enter after Tab will jump to that heading
- Interactive exercises require mouse/trackpad clicks
- Code editing requires clicking in code boxes first

WHAT WORKS WITH KEYBOARD:
- Ctrl+F to search within the tutorial
- Tab to navigate between section headings
- Screen readers can access all text content

IMPORTANT: This tutorial platform has limited keyboard accessibility.
Interactive elements like "Run Code" and "Submit Answer" buttons
require mouse/trackpad interaction.`);
    }
  });
  
  // Screen Reader Instructions - Collapsible accessibility guidance
  function addScreenReaderInstructions() {
    const instructionsDiv = document.createElement('div');
    instructionsDiv.setAttribute('role', 'region');
    instructionsDiv.setAttribute('aria-label', 'Accessibility Instructions');
    instructionsDiv.innerHTML = `
      <div style="background: #f8f9fa; border: 1px solid #dee2e6; border-radius: 6px; margin: 15px 0;">
        <button id="accessibility-toggle" style="width: 100%; padding: 12px; background: none; border: none; text-align: left; cursor: pointer; font-weight: bold; color: #495057;" 
                aria-expanded="false" aria-controls="accessibility-content">
          ▶ Accessibility & Navigation Guide (Click to expand)
        </button>
        
        <div id="accessibility-content" style="display: none; padding: 0 15px 15px 15px;">
          <div style="margin-bottom: 15px;">
            <h4 style="color: #6c757d; font-size: 16px; margin-bottom: 8px;">Keyboard Navigation:</h4>
            <ul style="margin: 0; padding-left: 20px; line-height: 1.6; font-size: 14px;">
              <li><strong>Press Ctrl+Shift+?</strong> for detailed keyboard help dialog</li>
              <li>Tab moves between section headings only</li>
              <li>Enter after Tab jumps to that heading</li>
              <li>Ctrl+F to search within tutorial</li>
            </ul>
          </div>
          
          <div style="margin-bottom: 15px;">
            <h4 style="color: #6c757d; font-size: 16px; margin-bottom: 8px;">For Screen Reader Users:</h4>
            <ul style="margin: 0; padding-left: 20px; line-height: 1.6; font-size: 14px;">
              <li>All text content is accessible to screen readers</li>
              <li>Interactive code exercises require mouse interaction</li>
              <li>Success messages will be announced as they appear</li>
            </ul>
          </div>
          
          <div style="margin-bottom: 15px;">
            <h4 style="color: #6c757d; font-size: 16px; margin-bottom: 8px;">Platform Limitations:</h4>
            <ul style="margin: 0; padding-left: 20px; line-height: 1.6; font-size: 14px;">
              <li>Exercise buttons require mouse clicks</li>
              <li>Code editing needs mouse interaction</li>
              <li>Interactive elements cannot be reached via Tab</li>
            </ul>
          </div>
          
          <div style="margin-bottom: 15px;">
            <h4 style="color: #6c757d; font-size: 16px; margin-bottom: 8px;">Recommended Approach:</h4>
            <ul style="margin: 0; padding-left: 20px; line-height: 1.6; font-size: 14px;">
              <li>Read content using screen reader navigation</li>
              <li>Use heading navigation between sections</li>
              <li>Consider working with assistance for interactive exercises</li>
            </ul>
          </div>
          
          <p style="margin: 0; font-size: 13px; color: #6c757d; font-style: italic;">
            This platform has inherent accessibility limitations. We've enhanced what's possible while being transparent about constraints.
          </p>
        </div>
      </div>
    `;
    
    // Add click handler for toggle
    setTimeout(function() {
      const toggleBtn = document.getElementById('accessibility-toggle');
      const content = document.getElementById('accessibility-content');
      
      if (toggleBtn && content) {
        toggleBtn.addEventListener('click', function() {
          const isExpanded = content.style.display !== 'none';
          content.style.display = isExpanded ? 'none' : 'block';
          toggleBtn.setAttribute('aria-expanded', !isExpanded);
          toggleBtn.innerHTML = (isExpanded ? '▶' : '▼') + ' Accessibility & Navigation Guide (Click to ' + (isExpanded ? 'expand' : 'collapse') + ')';
        });
      }
    }, 500);
    
    // Try multiple ways to find the right insertion point
    let insertLocation = null;
    
    // Try 1: Look for the h2 with "Introduction" text 
    const headings = document.querySelectorAll('h2');
    for (let i = 0; i < headings.length; i++) {
      if (headings[i].textContent.includes('Introduction')) {
        insertLocation = headings[i];
        break;
      }
    }
    
    // Try 2: Just use the first h2 if Introduction not found
    if (!insertLocation) {
      insertLocation = document.querySelector('h2');
    }
    
    // Insert the accessibility guide immediately after the heading
    if (insertLocation) {
      insertLocation.parentNode.insertBefore(instructionsDiv, insertLocation.nextSibling);
      console.log('Accessibility guide added successfully');
    } else {
      console.log('Could not find insertion point for accessibility guide');
    }
  }
  
  // Enhance success messages with exercise context
  setInterval(function() {
    const successMessages = document.querySelectorAll('.alert-success:not([data-enhanced])');
    successMessages.forEach(function(msg) {
      msg.setAttribute('data-enhanced', 'true');
      msg.setAttribute('role', 'status');
      msg.setAttribute('aria-live', 'polite');
      
      // Add exercise context to success messages
      const exerciseContainer = msg.closest('.tutorial-exercise');
      if (exerciseContainer) {
        const originalText = msg.textContent;
        if (!originalText.includes('Exercise completed')) {
          msg.textContent = 'Exercise completed successfully: ' + originalText;
        }
      }
    });
  }, 2000);
  
  // Enhance error messages for screen readers
setInterval(function() {
  const errorMessages = document.querySelectorAll('.alert-danger:not([data-enhanced]), .tutorial-exercise-error:not([data-enhanced])');
  errorMessages.forEach(function(msg) {
    msg.setAttribute('data-enhanced', 'true');
    msg.setAttribute('role', 'alert');
    msg.setAttribute('aria-live', 'assertive');
  });
}, 2000);

  // Run the accessibility guide function
  setTimeout(addScreenReaderInstructions, 1000);
  
});
</script>
```

## Learning Objectives

In this tutorial, you will learn how to describe an [**ordinal level variable**]{.important-text}, using summary statistics and plots. Specifically, we will cover:

* How to determine and describe the [**central tendency**]{.important-text} of an ordinal variable
* How to use the `median()` function
* How to summarize the [**dispersion**]{.important-text} of an ordinal variable

We will also practice the tools we used to generate frequency distributions and proportions for nominal variables using the functions in the package [dplyr]{.package-name}, and bar plots with the `ggplot()` function in the package [ggplot2]{.package-name}.

```{r setup, include=FALSE}
library(learnr)
library(tidyverse)
library(knitr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(error = TRUE)
counties <- qpaTutorials::counties
fHouse <- qpaTutorials::fHouse
sum(is.na(fHouse$FHStatus))
#fHouse <-filter(fHouse, FHStatus==)
fHouse$FHStatus2[fHouse$FHStatus=="NF"] <- 0
fHouse$FHStatus2[fHouse$FHStatus=="PF"] <- 1
fHouse$FHStatus2[fHouse$FHStatus=="F"] <- 2
fHouse$FHStatus2[fHouse$FHStatus=="NA"] <- NA

fHouse$FHStatus<-fHouse$FHStatus2
#fHouse$FHStatus <- factor(fHouse$FHStatus, ordered=TRUE, levels = c("NF", "PF", "F"))
```

## Univariate Description

As we discussed in the previous lesson, the first step in any analysis is to describe the variables that are the focus of our research question in order to convey how they behave, what they look like, and, therefore, what features of the data our explanation needs to account for. We introduced central tendency, dispersion, and the shape of the distribution as ways to summarize the information in a variable in the context of nominal variables. In this tutorial, we look at the tools available for describing ordinal variables.

**Practice: Identify ordinal level variables**
```{r concepts-quiz-ordinal, echo=FALSE}
question("Which of the following variables are ordinal level measures?",
  answer("The GDP measured in dollars of nations in the European Union", message="GDP is an interval (and ratio) measure where the differences between values are equal at all points."),
  answer("Freedom House measure of whether a country is free, partly free, or not free", correct = TRUE, message="The Freedom House measure is ordinal because the response options are ranked based on how free a country is. "),
  answer("The level of education of a respondent in a survey where the values include less than a high school degree, high school degree, college degree, and beyond college", correct = TRUE, message="Level of education is ordinal because each successive value indicates more education, but the differences between the categories are not the same. "),
  answer("A variable measuring whether a country supported a particular UN resolution", message = "The UN support variable is a nominal variable as the values taken denote unordered categories."),
  allow_retry = TRUE,
  try_again = "Hint: There are two ordinal variables in this list."
)
```

The methods we use for the descriptive analysis of ordinal variables are almost identical to those we use to describe nominal variables, so this tutorial will allow you to practice what you've learned as well as extend your knowledge.

## Central Tendency

We can calculate the central tendency of an ordinal variable using the mode, but, in addition, we can calculate the [**median**]{.important-text}.

### The mode

The data set, **counties**, contains the ordinal variable **rural_urban**. It is a 9-category scale measured by the US Department of Agriculture and coded as follows:

+ 1: Counties in metro areas of 1 million population or more
+ 2:	Counties in metro areas of 250,000 to 1 million population
+ 3:	Counties in metro areas of fewer than 250,000 population
+ 4:	Urban population of 20,000 or more, adjacent to a metro area
+ 5:	Urban population of 20,000 or more, not adjacent to a metro area
+ 6:	Urban population of 2,500 to 19,999, adjacent to a metro area
+ 7:	Urban population of 2,500 to 19,999, not adjacent to a metro area
+ 8:	Completely rural or less than 2,500 urban population, adjacent to a metro area
+ 9:	Completely rural or less than 2,500 urban population, not adjacent to a metro area.

Each category is more populous than the last, but the differences between categories are not identical.

We compute the mode of an ordinal variable in exactly the same fashion as for a nominal variable: we identify the most commonly occurring value from a frequency table or the table of proportions.  

See if you can create a frequency and proportion table for the variable **rural_urban** in the data frame **counties** using the [dplyr]{.package-name} functions. Name your frequency variable `freq` and your proportion variable `prop`.  Don't forget to load the package with the `library()` function!

**Practice: Create a frequency table for the rural variable**
```{r analysis-practice-rural-table, exercise = TRUE, exercise.lines=4, exercise.cap="Interactive exercise: Write code from scratch to create a frequency table for the rural variable. Click Submit Answer to check correctness. Four hints are available."}

```

```{r analysis-practice-rural-table-hint-1}
# Hint 1 of 4
Start by loading the dplyr package needed to use these functions.
```

```{r analysis-practice-rural-table-hint-2}
# Hint 2 of 4
Proceed by writing a four-step pipeline: begin with the data frame that contains
rural_urban, then use group_by(), summarise(), and mutate(), each on its own line and
connected with the pipe operator.
```

```{r analysis-practice-rural-table-hint-3}
# Hint 3 of 4
In the grouping step, pass rural_urban to group_by() so that the rows are grouped by rural
versus urban. In the next step, use summarise() to create a new variable called freq that
counts the number of rows in each rural_urban group using n().

```

```{r analysis-practice-rural-table-hint-4}
# Hint 4 of 4
After you have freq, add a mutate() step to create a proportion  variable called prop.
Inside mutate(), define prop as freq divided by the total of freq across both groups, 
using sum(freq) in the denominator.
```


```{r analysis-practice-rural-table-solution}
library(dplyr)
counties %>%
  group_by(rural_urban) %>%
  summarise(freq = n()) %>%
  mutate(prop = freq / sum(freq))
```

```{r analysis-practice-rural-table-check}
grade_code()
```

**Practice: Select a statement that correctly interprets the modal category**
```{r concepts-quiz-mode-interpretation, echo=FALSE}
question("Which of the following statements *best* describes the mode of the variable rural_urban?",
  answer("There is no meaningful modal category", message = "While many of the rural-urban categories are similarly represented in the data, one category has one third more observations than any other."),
  answer("The modal category is 6", message="It is better to provide substantive meaning than to list the numeric code."),
  answer("The modal category is an urban population of 2,500 to 19,999, adjacent to a metro area", correct = TRUE, message="The modal category for **rural_urban** is the 6th category, but it is more informative to provide a substantive interpretation of the mode. In other words, the typical county is one with an Urban population of 2,500 to 19,999, adjacent to a metro area."),
  allow_retry = TRUE
)
```




### The median

We can also use the [**median**]{.important-text} to calculate the central tendency of an ordinal variable.  The median value is the value in the middle when we list all values in rank order for all cases. Thus, it is the value at the 50th percentile: half the cases have values at or below the median and half have values at or above the median. If there are an even number of cases, we split the difference between the middle two values. 

Can you tell why the median is not an appropriate measure of central tendency for a nominal variable? There is no inherent ordering to nominal variables, so we cannot rank them to find the middle value. We can't rank racial categories or regions of the country as they have no intrinsic numeric values.  If a nominal variable is coded numerically, however, it is  possible -- **but always inappropriate ** -- to identify a median.  Don't do it! 

R comes with a simple to use function to calculate the median: `median()`!
Note that if the variable has any missing values we need to specify  `na.rm=TRUE`. Doing so removes the observations with missing values before calculating the median. Failing to do so, if there are missing values, will generate a median equal to NA.

**Try it: Calculate the median**
```{r analysis-calculate-median, exercise = TRUE, exercise.cap="Interactive exercise: Calculate the median. Use the Run Code button to execute and see the result."}
median(counties$rural_urban, na.rm = TRUE)
```

The median category for **rural_urban** is the same as the mode: 6. Half the counties are in the 6th category -- having an "Urban population of 2,500 to 19,999, adjacent to a metro area" -- or below and half the counties are in the 6th category or above. 



## Dispersion and the shape of the distribution

Like nominal level variables, we use the frequency distribution or the distribution of proportions in tabular or graphical form to describe the dispersion, or spread, of the variable and to describe the shape of the distribution.

Let's plot the frequency distribution and examine the dispersion from a bar plot.  Replace the XXXX in the following code to generate a bar plot with `ggplot()`.  Specify the title as "Number of US Counties that Cast a Majority of \\nVotes for the Major Party Presidential Candidates". Remember that the \\n tells `ggplot()` to add a line break to split the title into two lines. You will need to include it in the code below. Set the title for the x-axis label to `NULL`. This will omit the x-axis title. It's not needed here as having a title is redundant for this axis. 

**Practice: Create a bar chart for the rural_urban variable**
```{r visual-practice-bar-rural, exercise=TRUE, exercise.cap="Interactive exercise: Create a bar chart for the rural_urban variable. Fill in the XXXX placeholders and click Submit Answer to check correctness. Three hints are available."}
library(gpplot2)
ggplot(data = counties, XXXX = XXXX(x = factor(rural_urban))) +
         XXXX +
         labs(title = "XXXX", 
              subtitle = "2016",
              x = XXXX,  
              XXXX = "Frequency", 
              XXXX = "Source: Linn, Nagler, Zilinsky")
```

```{r visual-practice-bar-rural-hint-1}
# Hint 1 of 3
The first XXXX inside ggplot() should be the argument named `mapping`. That argument will
hold the aesthetic instructions for the plot.
```

```{r visual-practice-bar-rural-hint-2}
# Hint 2 of 3
The second XXXX is the function that wraps x = factor(rural_urban) to tell ggplot how to 
use that variable in the plot. Use the standard aesthetic-mapping function there.
```

```{r visual-practice-bar-rural-hint-3}
# Hint 3 of 3
Add the bar-chart layer before the labels. Then look at the remaining XXXX inside labs():
these are the argument names you must supply. One XXXX is the title argument (using the 
full text with the \n), one is the x-axis label (set this to NULL), and the last two XXXX
are the argument names for the y-axis label and the caption.
```


```{r visual-practice-bar-rural-solution}
library(ggplot2)
ggplot(data = counties, mapping = aes(x = factor(rural_urban))) +
         geom_bar() +
         labs(title = "Number of US Counties that Cast a Majority of \nVotes for the Major Party Presidential Candidates", 
              subtitle = "2016",
              x = NULL,  
              y = "Frequency", 
              caption = "Source: Linn, Nagler, Zilinsky")
```

```{r visual-practice-bar-rural-check}
grade_code()
```

Before we interpret the plot, let's make it more informative.

We'd like to replace the variable values 1-9 on the x-axis with informative phrases, but each category label is quite long and would not fit below the bars. We can solve this problem by flipping the coordinates. This is easy to do by adding another layer to the `ggplot()`.  The code below illustrates the `coord_flip()` function and replaces 1-9 with labels from the codebook entry.  Look carefully at the `scale_x_discrete()` function and how the labels are added. I've also filled the bars in the color `"green"` using the `fill` argument inside `geom_bar()`.

**Explore: Flip coordinates and add custom labels**
```{r visual-explore-flip-and-label, exercise=TRUE, exercise.cap="Interactive exercise: Flip coordinates and add custom labels. Run the code, then try removing coord_flip() or scale_x_discrete() to see what happens."}
ggplot(data = counties, mapping = aes(x = factor(rural_urban))) +
  geom_bar(fill = "green") +
  labs(title = "USDA Urban/Rural Status", 
       subtitle = "US Counties in 2016",
       x = NULL, 
       y = "Frequency",
       caption = "Source: Linn, Nagler, Zilinsky") +
  scale_x_discrete(labels = c("1 Million","250,000 to 1 million", "< 250,000", "20k+ near metro area", "20k+ not near  metro area", "2,500-19,999 near metro area", "2,500-19,999 not near metro area", "< 2,500 near metro area", "< 2,500 not near metro area", "NA")) +
  coord_flip()
```

The frequency distribution and bar plot illustrate the central tendency, dispersion, and shape of the  distribution of an ordinal variable in the same way as for a nominal variable. For this 9 category variable, we see:

1. The modal category is "Urban population of 2,500 to 19,999, adjacent to a metro area".
2. The number of counties in many categories is somewhat similar. The data is highly dispersed.
3. The distribution is somewhat uniform. This makes the mode less informative about the typical county’s rural-ness than if one category dominated the data, in which case the data would be unimodal. If half the counties were very urban and half very rural, the distribution would be bimodal. These distinctions are somewhat subjective, so we usually report the frequency (proportion) distribution or the bar plot and let the reader judge.

Note that we cannot easily determine the median from the plot.

### Your turn

Let’s look at another ordinal variable. We will examine the variable **FHStatus**, Freedom House's measure of whether a country is "Not Free" (coded 0), "Partly Free" (coded 1), or "Free" (coded 2). The variable is in the data set **fHouse**.  The variable is ordinal rather than categorical because as you go from "Not Free" to "Free," the country becomes *freer*. The variable constitutes a ranking of countries on this dimension.

First, we will generate a frequency and proportion table using the <span style="color:DarkGreen">dplyr</span> functions.  You do not need to load the package again because we've done so in an earlier exercise.  See if you can produce the table without looking at the hints. Name the counts variable **freq** and the proportion variable **prop**.

**Practice: Create a frequency and proportion table for FHStatus**
```{r analysis-practice-fhstatus-freq, exercise = TRUE, exercise.lines=4, exercise.cap="Interactive exercise: Write code from scratch to create a frequency and proportion table for the FHStatus variable. Name the new columns 'freq' and 'prop'. Click Submit Answer to check correctness. Five hints are available."}

```



```{r analysis-practice-fhstatus-freq-hint-1}
# Hint 1 of 5
Begin by starting your pipeline with the fHouse dataset so that every subsequent step 
in your code operates on the correct source of data.
```

```{r analysis-practice-fhstatus-freq-hint-2}
# Hint 2 of 5
Next, apply group_by() to FHStatus so the observations are organized by each category 
before you create any summaries or new variables.
```

```{r analysis-practice-fhstatus-freq-hint-3}
# Hint 3 of 5
Use summarise() to produce a column named freq, making sure to count the number of cases 
in each FHStatus group by calling n() inside it.
```


```{r analysis-practice-fhstatus-freq-hint-4}
# Hint 4 of 5
After computing frequencies, use mutate() to add a prop column defined as freq divided by
the total sum of freq across all grouped categories.
```

```{r analysis-practice-fhstatus-freq-hint-5}
# Hint 5 of 5
Keep each step connected with the %>% operator at the end of the line so the output from 
one function is passed directly into the next step.
```

```{r analysis-practice-fhstatus-freq-solution}
fHouse %>%
  group_by(FHStatus) %>%
  summarise(freq = n()) %>%
  mutate(prop = freq / sum(freq))
```

```{r analysis-practice-fhstatus-freq-check}
grade_code()
```

**Practice: Interpret the mode of FHStatus**
```{r concepts-quiz-fhstatus-mode, echo=FALSE}
question("Which of the following statements *best* describes the mode of the variable FHStatus?",
  answer("The mode is 2", message = "Rather than refer to a value as 2, we refer to the modal category. "),
  answer("The modal category is 2", message="It is better to provide substantive meaning than to list the numeric code. "),
  answer("The modal category is Free", correct = TRUE, message="The modal category for **FHStatus** is the 2nd category, but it is more informative to provide a substantive interpretation of the mode. In other words, the typical country is free."),
  answer("The distribution is multi-modal", message="One single category is more prevalent than all others,"),
  allow_retry = TRUE
)
```

Calculate the median.

**Practice: Calculate the median for FHStatus**
```{r analysis-practice-fhstatus-median, exercise=TRUE, exercise.cap="Interactive exercise: Write the code to calculate the median for the FHStatus variable. Remember to handle missing values. Click Submit Answer to check correctness. Three hints are available."}

```

```{r analysis-practice-fhstatus-median-solution}
median(fHouse$FHStatus, na.rm=TRUE)
```

```{r analysis-practice-fhstatus-median-hint-1}
# Hint 1 of 3
Start by using the median() function, which computes the middle value of a numeric or
ordinal variable. You will need to supply both the dataset and the specific variable 
inside its parentheses.
```

```{r analysis-practice-fhstatus-median-hint-2}
# Hint 2 of 3
When passing the variable to median(), use the $ notation to access FHStatus from the 
fHouse dataframe. Make sure the variable reference appears directly inside the median() call.
```

```{r analysis-practice-fhstatus-median-hint-3}
# Hint 3 of 3
Because the prompt mentions handling missing values, include the na.rm = TRUE argument
inside median() so that any NA entries are ignored when computing the median. This 
ensures the result is calculated correctly for this variable.
```

```{r analysis-practice-fhstatus-median-check}
grade_code()
```


A median of 1 means that the middle country in the distribution falls into the category coded as 1, which is “Partly Free.” In other words, if you lined up all countries from least free (0) to most free (2) and looked at the one in the exact center, that country’s status would be Partly Free. Note that the median of an ordinal variable depends on the order of the categories. If the coding were flipped—for example, If the ordering itself was reversed the median would change. This is why it is essential to state clearly how the categories are ordered whenever you report a median for an ordinal variable.

Let's plot the frequency distribution as a bar plot for practice. Try to generate the plot from scratch without looking at the hints. To produce an informative plot, use the `title` "Freedom House Ranking of Countries", label the `y` axis "Frequency", relabel the x-axis tick marks by setting "0" to "Not Free", "1" to "Partly Free", and "2" to "Free" in the `labels()` function so that the axis tick mark labels. Omit the missing values from your plot using `na.translate = FALSE` inside the `scale_x_discrete()` function. Do not include the `subtitle` or `caption` label. Since there is no default value for either, this will simply omit the subtitle and caption from the plot.

**Practice: Create a frequency bar plot for FHStatus**
```{r visual-practice-fhstatus-freq-bar, exercise=TRUE, exercise.cap="Interactive exercise: Write ggplot code from scratch to plot FHStatus. See the text prompt for specific instructions on titles, labels, and axis scales. Click Submit Answer to check correctness. Six hints are available."}

```

```{r visual-practice-fhstatus-freq-bar-hint-1}
# Hint 1 of 6
Begin by calling ggplot() and supplying data = fHouse. 
```

```{r visual-practice-fhstatus-freq-bar-hint-2}
# Hint 2 of 6
Inside ggplot(), include a mapping argument that uses aes() to specify the x-variable you
want to display as the categories in your bar chart. Be sure to wrap in the factor()
function to tell ggplot the variable is nominal. Finish the line with a + so you can add
the geometry on the following new line.
```

```{r visual-practice-fhstatus-freq-bar-hint-3}
# Hint 3 of 6
Add geom_bar() on its own new line to create the bars. Since you will be adding more 
layers after this, end the geom_bar() line with a plus sign (+).
```

```{r visual-practice-fhstatus-freq-bar-hint-4}
# Hint 4 of 6
Use labs() on the next line to include the required main title and y-axis label. Leave out
the subtitle and caption entirely. End this labs() line with a + because the scale
adjustment still needs to be added afterward.
```

```{r visual-practice-fhstatus-freq-bar-hint-5}
# Hint 5 of 6
Add scale_x_discrete() on a new line to control the x-axis. Inside it, set na.translate =
FALSE and provide a c() vector of descriptive labels to replace the numeric tick marks.
Watch your commas and parentheses carefully. Make sure to list the labels in order, from
the one coded zero to the one coded 2.
```

```{r visual-practice-fhstatus-freq-bar-hint-6}
# Hint 6 of 6
Confirm that every plot layer except the final one ends with a +, and that each major
component—mapping, geometry, labels, and scale—appears on its own separate line to keep 
the structure readable and correctly formatted.
```

```{r visual-practice-fhstatus-freq-bar-solution}
ggplot(data = fHouse,  mapping = aes(x = factor(FHStatus))) +
         geom_bar() +
         labs(title = "Freedom House Ranking of Countries", 
              x = NULL,  
              y = "Frequency") +
        scale_x_discrete(na.translate = FALSE, labels = c("Not Free", "Partly Free", "Free"))
```

```{r visual-practice-fhstatus-freq-bar-check}
grade_code()
```


Let's recreate this plot to display the y-axis in proportions. You only have to change two lines: the `geom_bar()` line needs a y aesthetic and the y axis title needs to change to "Proportion".

**Practice: Convert the plot to show proportions**
```{r visual-practice-fhstatus-prop-bar, exercise=TRUE, exercise.cap="Interactive exercise: Modify the plot to show proportions instead of frequencies. Fill in the XXXX placeholders and click Submit Answer to check correctness. Three hints are available."}
ggplot(data = fHouse,  mapping = aes(x = factor(FHStatus))) +
         geom_bar(XXXX(XXXX = XXXX(XXXX), XXXX = XXXX)) +
         labs(title = "Freedom House Ranking of Countries", 
              x = NULL,  
              y = "XXXX") +
         scale_x_discrete(na.translate = FALSE, labels = c("Not Free", "Partly Free", "Free"))
```

```{r visual-practice-fhstatus-prop-bar-hint-1}
# Hint 1 of 3
To turn counts into proportions within geom_bar(), you need a y aesthetic that uses
after_stat(prop). This goes inside an aes() call in the XXXX so that ggplot knows 
to compute proportions instead of raw frequencies.
```

```{r visual-practice-fhstatus-prop-bar-hint-2}
# Hint 2 of 3
Inside that aes(), one XXXX should map y to the after_stat(prop) expression, and another
XXXX should set a grouping so the proportions are computed over all bars together (use one 
for the group value). Watch the commas and match all parentheses carefully in the
geom_bar() line.
```

```{r visual-practice-fhstatus-prop-bar-hint-3}
# Hint 3 of 3
The remaining XXXX is in labs(), where the y-axis title still refers to frequencies. 
Update that quoted text so the y label correctly describes what is now being shown on the
axis: proportions rather than counts.
```

```{r visual-practice-fhstatus-prop-bar-solution}
ggplot(data = fHouse,  mapping = aes(x = factor(FHStatus))) +
         geom_bar(aes(y = after_stat(prop), group = 1)) +
         labs(title = "Freedom House Ranking of Countries", 
              x = NULL,  
              y = "Proportion") +
         scale_x_discrete(na.translate = FALSE, labels = c("Not Free", "Partly Free", "Free"))
```

```{r visual-practice-fhstatus-prop-bar-check}
grade_code()
```



## The Takeaways

Describing ordinal variables looks very much like describing a nominal variable. The only addition to our arsenal of descriptive tools is the median.  To summarize what we've learned:

1. There are 3 types of descriptive information to help you convey the main features of a single variable: [**central tendency**]{.important-text}, [**dispersion**]{.important-text}, and the [**shape of the distribution**]{.important-text}. Each piece of information complements the others.
2. The appropriate tools for description depend on the level of measurement of the variable you wish to describe. The information here describes what is appropriate for ordinal variables.

    a. The [**mode**]{.important-text} is the most commonly occurring category across cases and is the only statistic appropriate for describing the central tendency of a nominal variable. It can also be used to describe the central tendency of an ordinal variable. In addition, we can use the [**median**]{.important-text}. The median value marks the middle value or the value at the 50th percentile when all observations' values are rank-ordered. Both the mode and median may be more or less useful descriptions of the center of the distribution depending on the nature of the dispersion in the data, so often we present the entire frequency distribution for ordinal variables, just as for nominal variables.
    b. The frequency distribution, or the distribution of proportions, reveals the [**dispersion**]{.important-text} in the data, i.e., it tells us how many or what proportion of cases take each value of the variable.  It also makes clear which category is the modal value and is a numeric representation of the [**distribution**]{.important-text}.
    c. Bar plots are often used to visualize the central tendency, dispersion, and shape of the distribution of ordinal variables in the same manner as nominal variables, particularly if the specific counts or proportions are not of central interest.

Which of these pieces of information should you present when conducting descriptive analysis on an ordinal variable? The answer is the same as for a nominal variable. The choice depends on the main feature or features of the data and what you wish to emphasize. You may report the modal category or categories, the median, and the number or proportion of cases in the category, or you may choose to present the entire distribution, either in a table or a plot. 



![](images/UnivariateOrdinal.png){width="70%" fig.alt="Infographic titled ‘Ordinal Level Variable’ with four stacked blue boxes: Median (or Mode); Percentiles/IQR (dispersion of ranks); Ordered distribution (e.g., unimodal/bimodal and skew on ranks); Bar plot of ordered categories or cumulative proportions. A left sidebar lists Central Tendency, Dispersion, Distribution, and Visualization."}


<p role="status"
   aria-label="Great work! You have completed this tutorial and strengthened your understanding of ordinal variables."
   style="font-size: 28px; font-weight: 600; color: #2e7d32; text-align: center; margin-top: 50px;">
  <span aria-hidden="true">🌟 Excellent work! 🌟</span>
</p>



