---
title: "Univariate Description: Ordinal Variables"
tutorial:
  id: "05-Ordinal-Description"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Learn how to derscribe an ordinal variable: central tendency, dispersion, bar plot.
---

## Learning Objectives

In this tutorial, you will learn how to describe an ordinal level variable, using summary statistics and plots. Specifically, we will cover:

* How to determine and describe the central tendency of an ordinal variable
* How to use the `median()` function
* How to summarize the dispersion of an ordinal variable

We will also practice the tools we used to generate frequency distributions and proportions for nominal variables using the functions in the package <span style="color:DarkGreen">dplyr</span>, and bar plots with the `ggplot()` function in the package <span style="color:DarkGreen">ggplot2</span>.

```{r setup, include=FALSE}
library(learnr)
library(tidyverse)
library(knitr)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(error = TRUE)
counties <- qpaTutorials::counties
fHouse <- qpaTutorials::fHouse
sum(is.na(fHouse$FHStatus))
#fHouse <-filter(fHouse, FHStatus==)
fHouse$FHStatus2[fHouse$FHStatus=="NF"] <- 0
fHouse$FHStatus2[fHouse$FHStatus=="PF"] <- 1
fHouse$FHStatus2[fHouse$FHStatus=="F"] <- 2
fHouse$FHStatus2[fHouse$FHStatus=="NA"] <- NA

fHouse$FHStatus<-fHouse$FHStatus2
#fHouse$FHStatus <- factor(fHouse$FHStatus, ordered=TRUE, levels = c("NF", "PF", "F"))
```

## Univariate Description

As we discussed in the previous lesson, the first step in any analysis is to describe the variables that are the focus of our research question in order to convey how they behave, what they look like, and, therefore, what features of the data our explanation needs to account for. We introduced central tendency, dispersion, and the shape of the distribution as ways to summarize the information in a variable in the context of nominal variables. In this tutorial, we look at the tools available for describing ordinal variables.


```{r letter-a, echo=FALSE}
question("Which of the following variables are ordinal level measures?",
  answer("The GDP measured in dollars of nations in the European Union", message="GDP is an interval (and ratio) measure where the differences between values are equal at all points."),
  answer("Freedom House measure of whether a country is free, partly free, or not free", correct = TRUE, message="The Freedom House measure is ordinal because the response options are ranked based on how free a country is. "),
  answer("The level of education of a respondent in a survey where the values include less than a high school degree, high school degree, college degree, and beyond college", correct = TRUE, message="Level of education is ordinal because each successive value indicates more education, but the differences between the categories are not the same. "),
  answer("A variable measuring whether a country supported a particular UN resolution", message = "The UN support variable is a nominal variable as the values taken denote unordered categories."),
  allow_retry = TRUE,
  try_again = "Hint: There are two ordinal variables in this list."
)
```

The methods we use for the descriptive analysis of ordinal variables are almost identical to those we use to describe nominal variables, so this tutorial will allow you to practice what you've learned as well as extend your knowledge.

## Central Tendency

We can calculate the central tendency of an ordinal variable using the mode, but, in addition, we can calculate the median.

### The mode

The data set, **counties**, contains the ordinal variable **rural_urban**. It is a 9-category scale measured by the US Department of Agriculture and coded as follows:

+ 1: Counties in metro areas of 1 million population or more
+ 2:	Counties in metro areas of 250,000 to 1 million population
+ 3:	Counties in metro areas of fewer than 250,000 population
+ 4:	Urban population of 20,000 or more, adjacent to a metro area
+ 5:	Urban population of 20,000 or more, not adjacent to a metro area
+ 6:	Urban population of 2,500 to 19,999, adjacent to a metro area
+ 7:	Urban population of 2,500 to 19,999, not adjacent to a metro area
+ 8:	Completely rural or less than 2,500 urban population, adjacent to a metro area
+ 9:	Completely rural or less than 2,500 urban population, not adjacent to a metro area.


We compute the mode of an ordinal variable in exactly the same fashion as for a nominal variable: we identify the most commonly occurring value from a frequency table or the table of proportions.  

See if you can create a frequency and proportion table for the variable **rural_urban** in the data frame **counties** using the <span style="color:DarkGreen">dplyr</span> functions. Name your frequency variable `freq` and your proportion variable `prop`.  Don't forget to load the package with the `library()` function!


```{r ruraltable, exercise = TRUE, exercise.lines=4}

```

```{r ruraltable-hint-1}
Remember we start with then dataframe name 'and then' use 
group_by, summarise, and mutate in conjunction with the 
pipe operator.
```

```{r ruraltable-hint-2}
Have you passed the variable rural_urban to group_by?
```

```{r ruraltable-hint-3}
Have you calculated the frequencies using the summarise 
function and n()?
```

```{r ruraltable-hint-4}
Have you calculated the proportions using the mutate function to 
generate a new variable equal to the frequency variable divided 
by the sum of that variable?
```

```{r ruraltable-solution}
library(dplyr)
counties %>%
  group_by(rural_urban) %>%
  summarise(freq=n()) %>%
  mutate(prop=freq/sum(freq))
```

```{r ruraltable-check}
grade_code()
```


```{r letter-b, echo=FALSE}
question("Which of the following statements *best* describes the mode of the variable rural_urban?",
  answer("There is no meaningful modal category", message = "While many of the rural-urban categories are similarly represented in the data, one category has one third more observations than any other."),
  answer("The modal category is 6", message="It is better to provide substantive meaning than to list the numeric code."),
  answer("The modal category is an urban population of 2,500 to 19,999, adjacent to a metro area", correct = TRUE, message="The modal category for **rural_urban** is the 6th category, but it is more informative to provide a substantive interpretation of the mode. In other words, the typical county is one with an Urban population of 2,500 to 19,999, adjacent to a metro area."),
  allow_retry = TRUE
)
```




### The median

We can also use the <span style="color:blue">**median**</span> to calculate the central tendency of an ordinal variable.  The median value is the value in the middle when we list all values in rank order for all cases. Thus, it is the value at the 50th percentile: half the cases have values at or below the median and half have values at or above the median. If there are an even number of cases, we split the difference between the middle two values. 

Can you tell why the median is not an appropriate measure of central tendency for a nominal variable? There is no inherent ordering to nominal variables, so we cannot rank them to find the middle value. We can't rank racial categories or regions of the country as they have no intrinsic numeric values.  If a nominal variable is coded numerically, however, it is  possible -- **but always inappropriate ** -- to identify a median.  Don't do it! 

R comes with a simple to use function to calculate the median: `median()`!
Note that if the variable has any missing values we need to specify  `na.rm=TRUE`. Doing so removes the observations with missing values before calculating the median. Failing to do so, if there are missing values, will generate a median equal to NA.

```{r median, exercise = TRUE}
median(counties$rural_urban, na.rm = TRUE)
```

The median category for **rural_urban** is the same as the mode: 6. Half the counties are in the 6th category -- having an "Urban population of 2,500 to 19,999, adjacent to a metro area" -- or below and half the counties are in the 6th category or above. 



## Dispersion and the shape of the distribution

Like nominal level variables, we use the frequency distribution or the distribution of proportions in tabular or graphical form to describe the dispersion, or spread, of the variable and to describe the shape of the distribution.

Let's plot the frequency distribution and examine the dispersion from a bar plot.  Replace the XXXX in the following code to generate a bar plot with `ggplot()`.  Specify the title as "Number of US Counties that Cast a Majority of \\nVotes for the Major Party Presidential Candidates". Remember that the \\n tells `ggplot()` to add a line break to split the title into two lines. You will need to include it in the code below. Set the title for the x-axis label to `NuLL`. This will omit the x-axis title. It's not needed here as having a title is redundant for this axis.


```{r ruralplot, exercise=TRUE}
ggplot(data=counties, XXXX=XXXX(x=factor(rural_urban))) +
         XXXX +
         labs(title = "XXXX", 
              subtitle = "2016",
              x=XXXX,  
              XXXX="Frequency", 
              XXXX="Source: Linn, Nagler, Zilinsky")
```


```{r ruralplot-hint-1}
The first code that needs replacing tells ggplot() how to map 
the variables to the x axis. Specify the mapping 
argument and the aes() function
```

```{r ruralplot-hint-2}
Specify the geom type in the second line of code.
```

```{r ruralplot-hint-3}
The remaining XXXX refer to labels. Think about what
information you want to appear for which labels and 
supply the correct argument.
```


```{r ruralplot-solution}
ggplot(data=counties, mapping=aes(x=factor(rural_urban))) +
         geom_bar() +
         labs(title = "Number of US Counties that Cast a Majority of \nVotes for the Major Party Presidential Candidates", 
              subtitle = "2016",
              x=NULL,  
              y="Frequency", 
              caption="Source: Linn, Nagler, Zilinsky")
```

```{r ruralplot-check}
grade_code()
```

Before we interpret the plot, let's make it more informative.

We'd like to replace the variable values 1-9 on the x-axis with informative phrases, but each category label is quite long and would not fit below the bars. We can solve this problem by flipping the coordinates. This is easy to do by adding another layer to the `ggplot()`.  The code below illustrates the `coord_flip()` function and replaces 1-9 with labels from the codebook entry.  I've also filled the bars in the color `"green"` using the `fill` argument inside `geom_bar()`.

```{r flippedplot, exercise=TRUE}
ggplot(data=counties, mapping=aes(x=factor(rural_urban))) +
  geom_bar(fill="green") +
  labs(title="USDA Urban/Rural Status", 
       subtitle = "US Counties in 2016",
       x=NULL, 
       y="Frequency",
       caption="Source: Linn, Nagler, Zilinsky") +
  scale_x_discrete(labels = c("1 Million","250,000 to 1 million","< 250,000", "20k+ near metro area", "20k+ not near  metro area", "2,500-19,999 near metro area", "2,500-19,999 not near metro area", "< 2,500 near metro area", "< 2,500 not near metro area", "NA")) +
  coord_flip()
```

The frequency distribution and bar plot illustrate the central tendency, dispersion, and shape of the  distribution of an ordinal variable in the same way as for a nominal variable. For this 9 category variable, we see:

1. The modal category is "Urban population of 2,500 to 19,999, adjacent to a metro area".
2. The number of counties in many categories is somewhat similar. The data is highly dispersed.
3. The distribution is somewhat uniform. This makes the mode less informative about the typical county’s rural-ness than if one category dominated the data, in which case the data would be unimodal. If half the counties were very urban and half very rural, the distribution would be bimodal. These distinctions are somewhat subjective, so we usually report the frequency (proportion) distribution or the bar plot and let the reader judge.

Note that we cannot easily determine the median from the plot.

### Your turn

Let’s look at another ordinal variable. We will examine the variable **FHStatus**, Freedom House's measure of whether a country is "Not Free" (coded 0), "Partly Free" (coded 1), or "Free" (coded 2). The variable is in the data set **fHouse**.  The variable is ordinal rather than categorical because as you go from "Not Free" to "Free," the country becomes *freer*. The variable constitutes a ranking of countries on this dimension.

First, we will generate a frequency and proportion table using the <span style="color:DarkGreen">dplyr</span> functions.  You do not need to load the package again because we've done so in an earlier exercise.  See if you can produce the table without looking at the hints. Name the counts variable **freq** and the proportion variable **prop**.


```{r FHtable, exercise = TRUE, exercise.lines=4}

```



```{r FHtable-hint-1}
Remember to use group_by, summarise, and mutate in conjunction 
with the pipe operator.
```

```{r FHtable-hint-2}
Did you start with the dataframe fHouse?
```

```{r FHtable-hint-3}
Have you passed the variable FHStatus to group_by?
Have you calculated the frequencies using the summarise 
function and n()?
```


```{r FHtable-hint-4}
Have you calculated the proportions using the mutate function 
to generate a new variable equal to the frequency variable 
divided by the sum of that variable?
```

```{r FHtable-hint-5}
Don't forget the pipe, %>%, needs to be used at the end
of each line telling the code "and then" execute the function
on the next line.
```

```{r FHtable-solution}
fHouse %>%
  group_by(FHStatus) %>%
  summarise(freq=n()) %>%
  mutate(prop=freq/sum(freq))
```

```{r FHtable-check}
grade_code()
```


```{r letter-3, echo=FALSE}
question("Which of the following statements *best* describes the mode of the variable FHStatus?",
  answer("The mode is 2", message = "Rather than refer to a value as 2, we refer to the modal category. "),
  answer("The modal category is 2", message="It is better to provide substantive meaning than to list the numeric code. "),
  answer("The modal category is Free", correct = TRUE, message="The modal category for **FHStatus** is the 2nd category, but it is more informative to provide a substantive interpretation of the mode. In other words, the typical country is free."),
  answer("The distribution is multi-modal", message="One single category is more prevalent than all others,"),
  allow_retry = TRUE
)
```

Calculate the median.

```{r FHStatusMed, exercise=TRUE}

```

```{r FHStatusMed-solution}
median(fHouse$FHStatus, na.rm=TRUE)
```

```{r FHStatusMed-check}
grade_code()
```



Most nations are "Free" (coded 2) according to Freedom House, and a half or less are "Partly Free" (coded 1) and another half or more are "Partly Free".  Note that if the variable had been recoded so that "Free" was coded 0 and "Not Free" was coded 2, the median category would have been "Free." The moral of the story is that you need to clearly state the order of the ranking (from more to less free or from less free to more free) when you provide a median.


Let's plot the frequency distribution as a bar plot for practice. Try to generate the plot from scratch without looking at the hints. To produce an informative plot, use the `title` "Freedom House Ranking of Countries", label the `y` axis "Frequency", relabel the x-axis tick marks by setting "0" to "Not Free", "1" to "Partly Free", and "2" to "Free" in the `labels()` function so that the axis tick mark labels. Omit the missing values from your plot using `na.translate=FALSE` inside the `scale_x_discrete()` function. Do not include the `subtitle` or `caption` label. Since there is no default value for either, this will simply omit the subtitle and caption from the plot.


```{r FHplot, exercise=TRUE}

```

```{r FHplot-hint-1}
Did you call the ggplot() function and provide the data?
```

```{r FHplot-hint-2}
Did you use the mapping argument to specify the aesthetics?
```

```{r FHplot-hint-3}
Did you use a plus sign and add the geom_bar() layer?
```

```{r FHplot-hint-4}
Did you remember that to replace the labels you need to 
put the new names in quotations and to separate
them with commas and place them inside c()?
```

```{r FHplot-hint-5}
Are all your parentheses matched?
Did you include a plus sign for each layer at the end of the line?
```

```{r FHplot-solution}
ggplot(data=fHouse,  mapping=aes(x=factor(FHStatus))) +
         geom_bar() +
         labs(title = "Freedom House Ranking of Countries", 
              x=NULL,  
              y="Frequency") +
        scale_x_discrete(na.translate=FALSE, labels=c("Not Free", "Partly Free", "Free"))
```

```{r FHplot-check}
grade_code()
```

## The Takeaways

Describing ordinal variables looks very much like describing a nominal variable. The only addition to our arsenal of descriptive tools is the median.  To summarize what we've learned:

1. There are 3 types of descriptive information to help you convey the main features of a single variable: <span style="color:blue">**central tendency**</span>, <span style="color:blue">**dispersion**</span>, and the <span style="color:blue">**shape of the distribution**</span>. Each piece of information complements the others.
2. The appropriate tools for description depend on the level of measurement of the variable you wish to describe. The information here describes what is appropriate for ordinal variables.

    a. The <span style="color:blue">**mode**</span> is the most commonly occurring category across cases and is the only statistic appropriate for describing the central tendency of a nominal variable. It can also be used to describe the central tendency of an ordinal variable. In addition, we can use the <span style="color:blue">**median**</span>. The median value marks the middle value or the value at the 50th percentile when all observations' values are rank-ordered. Both the mode and median may be more or less useful descriptions of the center of the distribution depending on the nature of the dispersion in the data, so often we present the entire frequency distribution for ordinal variables, just as for nominal variables.
    b. The frequency distribution, or the distribution of proportions, reveals the <span style="color:blue">**dispersion**</span> in the data, i.e., it tells us how many or what proportion of cases take each value of the variable.  It also makes clear which category is the modal value and is a numeric representation of the <span style="color:blue">**distribution**</span>.
    c. Bar plots are often used to visualize the central tendency, dispersion, and shape of the distribution of ordinal variables in the same manner as nominal variables, particularly if the specific counts or proportions are not of central interest.

Which of these pieces of information should you present when conducting descriptive analysis on an ordinal variable? The answer is the same as for a nominal variable. The choice depends on the main feature or features of the data and what you wish to emphasize. You may report the modal category or categories, the median, and the number or proportion of cases in the category, or you may choose to present the entire distribution, either in a table or a plot. 



![](images/UnivariateOrdinal.png)



