---
title: "Bivariate Description: Two Interval Variables"
tutorial:
  id: "09-Bivariate-Description-Interval-Data"
output: 
  qpaTutorials::qpa_tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Learn how to describe the relationship between two interval-level variables.
---


## Learning Objectives

In this tutorial, you will continue learning how to describe the relationship between two variables. This time we will focus on how to do so when both variables are measured at the interval level. Specifically, we will cover:

* How to conduct a descriptive analysis of the relationship between interval variables with scatter plots using the `ggplot()` function in the R package [ggplot2]{.package-name}.
* How to conduct a descriptive analysis of the relationship between two interval variables using Pearson's correlation coefficient and plots.



```{r setup, include=FALSE}
library(learnr)
library(tidyverse)
library(knitr)
library(poliscidata)
library(descr)
library(ggplot2)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(error = TRUE)
data("counties", package = "qpaTutorials")
counties$employ_pop_ratio_25_64 <- counties$employ_pop_ratio_25_64*100
counties$dem2p_percent <- counties$dem2p_vote_share*100
counties$rural <-1
counties$rural[counties$rural_urban<=3] <- 0
counties$rural[counties$rural_urban>=8] <-2

data("fHouse", package = "qpaTutorials")
sum(is.na(fHouse$FHStatus))
#fHouse <-filter(fHouse, FHStatus==)
fHouse$FHStatus2[fHouse$FHStatus=="NF"] <- 0
fHouse$FHStatus2[fHouse$FHStatus=="PF"] <- 1
fHouse$FHStatus2[fHouse$FHStatus=="F"] <- 2
fHouse$FHStatus2[fHouse$FHStatus=="NA"] <- NA

fHouse$FHStatus<-fHouse$FHStatus2
#fHouse$FHStatus <- factor(fHouse$FHStatus, ordered=TRUE, levels = c("NF", "PF", "F"))
world <- world
```



## Bivariate Description

In the previous tutorials on bivariate description, we considered how to describe the relationships between variables measured at the nominal and/or ordinal level. We presented two tools available to summarize these types of relationships. We produced and interpreted the information in a cross-tab, which is a table that shows the distribution of cases (counts or proportions) across values of the dependent variable for cases that have different values on the independent variable. We then showed how to visualize the information in the cross-tab using a grouped bar plot. 

In the last tutorial, we considered how to conduct a descriptive analysis of the relationship between an interval variable and a nominal/categorical (or ordinal) variable. Typically we either compare a single summary statistic, the mean, or we compare the distribution of the interval variable for each of the values of the nominal or ordinal variable. 

If our interest lies in the relationship between two interval-level variables, we can summarize the relationship between them using Pearson's correlation coefficient. The  [**correlation coefficient quantifies the strength and direction of the linear association between two interval-level variables**]{.important-text}.  The logic of correlation is most easily understood in the context of a scatter plot. So we will introduce them first. A scatter plot provides a visual display of the nature, degree, and strength of the relationship between two interval-level variables. It presents pairs of values for each unit of analysis on two variables measured at the interval level. The position of each dot on the horizontal and vertical axis indicates values for an individual data point.


Let's start with a picture of a scatter plot:

```{r scatterplot-clintonvote-by-education, echo=FALSE, out.width='70%', fig.alt="Scatter plot showing the relationship between proportion of college graduates (x-axis) and Clinton vote share (y-axis) across U.S. counties in 2016, with blue points and reference lines marking average values. The points show a positive trend, where counties with higher levels of education tend to have higher Clinton vote shares, though there is substantial variation. The plot illustrates how scatterplots reveal relationships and variability between two continuous variables."}
knitr::include_graphics("images/samplescattervoteed.png")
```


This scatter plot shows the relationship between the two interval-level measures **dem2p_vote_share**, the share of the presidential vote cast for Hillary Clinton in 2016, and **prop_college_grad**, the proportion of college graduates in US counties. Each dot represents one unit of analysis, i.e., one county. Its placement on the x-axis represents its value on the variable plotted on the x-axis, i.e., proportion college graduates. It's placement on the y-axis represents its value on the variable plotted on the y-axis, i.e., proportion voting for Clinton. The horizontal (vertical) line is located at the mean value of the proportion voting for Clinton (proportion college graduates).

Consider the point in the upper right and labeled Arlington, VA. This point corresponds to 0.74, the share of adults who graduated college, (plotted on the x-axis) and a 0.82 share of the votes cast for Clinton (plotted on the y-axis).


```{r scatterplot-clintonvote-by-education-label-arlington, echo=FALSE, out.width='70%', fig.alt="Scatter plot showing the relationship between proportion of college graduates (x-axis) and Clinton vote share (y-axis) across U.S. counties in 2016, with blue points and horizontal and vertical lines marking the means. Many points cluster in the lower-left and upper-right quadrants, indicating that counties with lower education tend to have lower Clinton vote share and those with higher education tend to have higher Clinton vote share; an example high-education, high-vote-share county (Arlington, VA) is labeled. The plot demonstrates how quadrant patterns around mean lines reveal a positive relationship between two variables."}
knitr::include_graphics("images/samplescattervoteedArlington.png")
```

Points -- like Arlington county VA -- in the upper right quadrant have values of **prop_college_grads** above the mean and values of **dem2p_vote_share** above the mean. Points in the lower left quadrant have values of **prop_college_grads** below the mean and values of **dem2p_vote_share** below the mean. When points are clustered in these two quadrants -- as in this example -- there is a positive relationship between the two variables.

Scatter plots are widely used and very informative about the nature of the relationship between two interval-level variables. We can use the summary statistic, Pearson's correlation coefficient, to quantify the strength and direction of the relationship. Pearson's correlation coefficient, denoted  $r$, is given by the following formula:

$$r=\frac{1}{n-1} \Sigma \frac{ (x_i-\bar{x}) (y_i-\bar{y}) }{s_x s_y}$$
where

+ $x_i,y_i$ are the individual unit observations of $x$ and $y$, respectively
+ $\bar{x},\bar{y}$ are the sample means of $x$ and $y$, respectively
+ $s_x,s_y$ are the sample standard deviations of $x$ and $y$, respectively
+ $n$ is the sample size

Think through what this formula tells us.  When unit $i$'s value of both $x$ and $y$ are  above (or below) $\bar{x}$ and $\bar{y}$, $i$'s contribution to $r$ is positive. The more often that is true in our data, the greater $r$. In contrast, when  unit $i$'s  value is greater than the mean on one variable and less than the mean on the other, $i$'s contribution to $r$ is negative. The more often that is true in our data, the more $r$ will be negative.  If there is no systematic relationship between $x$ and $y$ then $r=0$.

The value of $r$ ranges between -1 and +1. 

+  $r>0$ indicates a positive relationship between $x$ and $y$: as one gets larger, the other gets larger.
+ $<0$ indicates a negative relationship: as one gets larger, the other gets smaller.
+ $r=0$ indicates no linear relationship.

Below is a picture illustrating the relationship between the pattern in a scatter plot and the correlation coefficient, $r$. A line is drawn through the points such that it is as close to all the points as possible. This helps us to gauge the relationship.  If the line is downward sloping (goes from upper left to lower right), as on the left, the correlation will be negative: higher values of one variable are associated with lower values of the other variable. If the line is upward sloping (goes from lower left to upper right), as on the right, the correlation will be positive: higher values of one variable are associated with higher values of the other variable. Finally, when the line is flat, as in the plot in the middle, there is no correlation between the two variables.  The steepness of the line is positively related to the strength of the correlation.


```{r scatterplot-correlation, echo=FALSE, out.width='70%', fig.alt="Three side-by-side scatter plots illustrating correlation between two variables with fitted lines: the left plot shows a strong negative relationship (r ≈ -0.86) with a downward-sloping line, the middle plot shows no relationship (r ≈ 0) with a flat line and scattered points, and the right plot shows a strong positive relationship (r ≈ 0.87) with an upward-sloping line. The figure demonstrates how the direction and strength of a linear pattern in a scatter plot correspond to the sign and magnitude of the correlation coefficient."}
knitr::include_graphics("images/Samplecorrplots.png")
```


## Scatter plots

We follow the same basic steps to produce a scatter plot using `ggplot()` in the R package [ggplot2]{.package-name} as we have previous plots.

1.  Tell `ggplot()` the data that contains the variable(s) to plot.  
2.  Tell `ggplot()` which variable(s) to plot and on which axis. We do this by mapping the variables to the appropriate axis using the `mapping` argument and the `aes()` function. This time we need to specify a variable for each the x and y axis. By convention, the outcome variable is on the y-axis.
3.  Tell `ggplot()` how to plot the data by adding a **geom** layer. To create a scatter plot we use `geom_point()`.
4.  Add layers that make the plot more informative by adding labels and formatting scales.


To illustrate how to make a scatter plot we will ask the question how solid is support for democracy around the world? The data set **fHouse** contains the variables **Pro.Dem** and **FHTotalAggr**. **ProDem** is coded as the percent of respondents in a country that said "very good" or "fairly good" in response to the following question asked in the World Values Survey (Wave 5, 2005-2008) and the European Values Survey (2008):

>“I’m going to describe various types of political systems and ask what you think about each as a way of governing this country. For each one, would you say it is a very good, fairly good, fairly bad or very bad way of governing this country? ...Having a democratic system?

**FHTotalAggr** is the Freedom House (2017) measure of a country's composite level of democracy based on civil liberties and political rights, Freedom House, 2017. Higher scores are more democratic. 

We might hypothesize the following:

::: {.hypothesis}
In a comparison of [countries], those having [a higher level of democracy] will be more likely to [have higher support for democratic government] than will those having [a lower level of democracy].
:::

I've loaded the [ggplot2]{.package-name} package for you. Try to make a scatter plot based on the instructions above.

**Practice: Create scatter plot of democracy and support**
```{r exercise-scatterplot-prodemocracy-and-democracylevel, exercise=TRUE, exercise.cap="Interactive exercise: Write code to create a scatter plot of ProDem by FHTotalAggr to examine the relationship between level of democracy and support for democratic government. Use the Run Code button to execute and click Submit Answer to check correctness. Two hints are available.", fig.alt="Scatter plot showing support for democracy (Pro.Dem) on the x-axis and level of democracy (FHTotalAggr) on the y-axis across countries. The points are widely scattered with no clear upward or downward trend, indicating little to no strong linear relationship between public support for democracy and institutional democracy levels. The plot demonstrates how scatter plots can be used to assess the presence or absence of a relationship between two variables."}
library(ggplot2)
 
```

```{r exercise-scatterplot-prodemocracy-and-democracylevel-hint-1}
# Hint 1 of 2
A scatterplot is used to examine the relationship between two numeric 
variables. Here, you are looking at how pro-democracy attitudes relate to 
the level of democracy. The explanatory variable should go on the x-axis 
and the outcome variable on the y-axis, so decide which variable plays 
each role before writing any code.
```

```{r exercise-scatterplot-prodemocracy-and-democracylevel-hint-2}
# Hint 2 of 2
Start with ggplot() using the fHouse data.

Inside aes(), map your chosen x and y variables.

Then use + to add a layer with geom_point() 
to display the individual observations.
```


```{r exercise-scatterplot-prodemocracy-and-democracylevel-solution}
library(ggplot2)
ggplot(data = fHouse, mapping = aes(x = Pro.Dem, y = FHTotalAggr)) +
  geom_point() 
```

```{r exercise-scatterplot-prodemocracy-and-democracylevel-check}
grade_code()
```

Before we interpret our plot, let's clean it up!  The output above provides a warning that it has "Removed 125 rows containing missing values (geom_point)."  We can avoid this by filtering out the missing values using the `filter()` function in the [dplyr]{.package-name} package and piping the result to `ggplot()`. 


I'm going to make the points purple by specifying `color = "purple"` inside the `geom_point()` function. Try editing the plot to pick a new color (For a list of the color names available see http://www.stat.columbia.edu/~tzheng/files/Rcolor.pdf)

I'm also going to add a `geom_smooth()` layer to the plot (remember the plus sign). This adds a line through the points and aids in our interpretation. I'm making the line orange using the `color` argument inside `geom_smooth()`, telling ggplot not to add standard errors around the line (`se = FALSE`), and to use the method of least squares to draw the line (`method = "lm"`). We will talk more about least squares lines in the coming classes.

We need to load the [dplyr]{.package-name} package with the `library()` function.

Look carefully at the code to make sure you understand what each line of the code is doing. You will be asked to make your own scatter plots in class.

**Try it: Customize scatter plot with color and trend line**

```{r scatterplot-prodemocracy-democracy-level-smooth, exercise=TRUE, exercise.cap="Interactive exercise: Run the code to customize a scatter plot by changing point color and adding a fitted trend line to visualize the relationship between Pro.Dem and FHTotalAggr. Use the Run Code button to execute and see results. No hints are available.", fig.alt="Scatter plot showing support for democratic rule (x-axis) and level of democracy (y-axis) across countries, with purple points and an orange fitted regression line. The points are widely dispersed with no strong pattern, and the fitted line is nearly flat with a slight negative slope, indicating little to no linear relationship between public support and institutional democracy. The plot demonstrates how adding a regression line helps assess the direction and strength of a relationship in a scatter plot."}
library(dplyr)
fHouse %>%
  filter(!is.na(FHTotalAggr), !is.na(Pro.Dem)) %>%
  ggplot(mapping = aes(x=Pro.Dem, y = FHTotalAggr)) +
  geom_point(color = "purple") +
  geom_smooth(color = "orange", se = FALSE, method = "lm") +
  labs(title = "Support for Democratic Values \nand Level of Democracy", 
       x = "Support for Democratic Rule", 
       y = "Level of Democracy", 
       caption = "Source: WVS and EVS, 2005-2008, \nFreedom House 2017") 
```


We can see that the points are quite widely scattered. In other words, there does not appear to be a strong relationship between the two variables. The line through the points is downward sloping but fairly flat, which also points to a weak (negative) relationship. There is no support for our hypothesis.

For more on scatter plots with `ggplot()` see: http://t-redactyl.io/blog/2016/02/creating-plots-in-r-using-ggplot2-part-5-scatterplots.html.



## The `cor()` function

Let's quantify the strength and direction of the linear relationship by calculating Pearson's correlation coefficient.

The `cor()` function in base R will accomplish this task. We will pass the function 4 arguments. The first two arguments are `x` and `y`, the two interval-level variables for which we desire the correlation coefficient. In addition, we need to include the `use` argument if there are any missing values. It tells R how to handle them. The default method ("everything") will return NA if there are missing values in any variables. We will set `use = "pairwise.complete.obs"`. This means the function will include a unit (country in our example) if both units have a value on both variables. Finally, we need to tell the function the specific correlation measure to compute. We do so by setting the `method` argument equal to "pearson".

**Try it: Calculate Pearson correlation between two variables**

```{r correlation-prodemocracy-democracylevel, exercise=TRUE, exercise.cap="Interactive exercise: Write code to calculate the Pearson correlation between Pro.Dem and FHTotalAggr, handling missing values appropriately. Use the Run Code button to execute and click Submit Answer to check correctness. No hints are available."}
cor(x = fHouse$Pro.Dem, y = fHouse$FHTotalAggr, use = "pairwise.complete.obs", method = "pearson")
```

The result, $r=-0.04$, is negative and close to zero. Thus we conclude that the relationship between these two variables is negative and weak. Perhaps surprisingly, there appears to be next to no relationship between support for democracy and level of democracy.

The data frame **fHouse** also contains the variable **Pro.Strong.Leader**. This is an interval variable that records the percent of respondents in a country that said "very good" or "fairly good" in response to the following question asked in the World Values Survey (Wave 5, 2005-2008) and the European Values Survey (2008):

> “I’m going to describe various types of political systems and ask you about each one as a way of governing the country. For each one, would you say it is a very good, fairly good, fairly bad or very bad way of governing this country?” ... “Having a strong leader who does not have to bother with parliament and elections?”


We might hypothesize the following:

::: {.hypothesis}
In a comparison of [countries], those having [a higher level of democracy] will be more likely to [have lower support for a strong leader] than will those having [a lower level of democracy].
:::

Edit the code below to produce a scatter plot illustrating the relationship between this variable and level of democracy (**FHTotalAggr**) from our previous example. We will begin by filtering out the missing values.  Try making the points blue and adding a red line.

**Explore: Plot democracy vs support for strong leader**

```{r exercise-scatterplot-democracylevel-strongleader, exercise=TRUE, exercise.cap="Interactive exercise: Fill in the XXXX values to create a scatter plot of Pro.Strong.Leader and FHTotalAggr, including filtering missing values and customizing point and line colors. Click Submit Answer to check correctness. Three hints are available.", fig.alt="Scatter plot showing support for a strong leader (x-axis) and level of democracy (y-axis) across countries, with blue points and a red fitted regression line. The points show a clear downward trend, with higher support for a strong leader associated with lower levels of democracy, as reflected in the negatively sloped line. The plot demonstrates how scatter plots and regression lines can reveal a negative relationship between two variables consistent with the hypothesis."}
fHouse %>%
  filter(XXXX(FHTotalAggr), XXXX(Pro.Strong.Leader)) %>%
ggplot(mapping = aes(x = Pro.Strong.Leader, y = FHTotalAggr)) +
 geom_point(XXXX) +
  geom_smooth(XXXX) +
  labs(title = "Support for Democratic Values \nand Level of Democracy",
       x = "Support for a Strong Leader", 
       y = "Level of Democracy", 
       caption = "Source: WVS and EVS, 2005-2008, \nFreedom House 2017")
```

```{r exercise-scatterplot-democracylevel-strongleader-hint-1}
# Hint 1 of 3
You are examining the relationship between support for a strong leader 
and level of democracy. Before plotting, make sure the data you pass 
through the pipe does not include missing values for either variable, 
using filter() to keep only cases where both are observed.
```

```{r exercise-scatterplot-democracylevel-strongleader-hint-2}
# Hint 2 of 3
After filtering, pipe the data into ggplot().

Inside aes(), map support for a strong leader to x and 
level of democracy to y. Then use + to add geom_point() 
to display the observations.
```

```{r exercise-scatterplot-democracylevel-strongleader-hint-3}
# Hint 3 of 3
Add a fitted line using geom_smooth().

Set se = FALSE and method = "lm", and specify colors inside 
geom_point() and geom_smooth() to customize the appearance.
```


```{r exercise-scatterplot-democracylevel-strongleader-solution}
fHouse %>%
  filter(!is.na(FHTotalAggr), !is.na(Pro.Strong.Leader)) %>%
ggplot(mapping = aes(x = Pro.Strong.Leader, y = FHTotalAggr)) +
 geom_point(color = "blue") +
  geom_smooth(color = "red",se = FALSE, method = "lm") +
  labs(title = "Support for Democratic Values \nand Level of Democracy",
       x = "Support for a Strong Leader", 
       y = "Level of Democracy", 
       caption = "Source: WVS and EVS, 2005-2008, \nFreedom House 2017")
```

```{r exercise-scatterplot-democracylevel-strongleader-check}
grade_code()
```

We can see the relationship is negative. For countries in which there is more support of a strong leader there  tends to be less support for democracy.

Edit the code below to estimate Pearson's correlation coefficient.

**Explore: Calculate correlation for strong leader and democracy**

```{r exercise-correlation-democracylevel-strongleader, exercise=TRUE, exercise.cap="Interactive exercise: Fill in the XXXX values to calculate the Pearson correlation between Pro.Strong.Leader and FHTotalAggr, handling missing values appropriately. Click Submit Answer to check correctness. One hint is available."}
cor(XXX = XXX, y = XXX$XXX, use = XXX, method = "XXX")
```

```{r exercise-correlation-democracylevel-strongleader-hint}

# Hint 1 of 1
The x and y arguments in cor() are just the two variables whose 
relationship you want to measure. Unlike a scatter plot, it does 
not matter which variable goes in x and which goes in y.

Use the two variables named in the question, written as fHouse$variable.

Set method = "pearson" and choose the use option that handles missing values.
```

```{r exercise-correlation-democracylevel-strongleader-solution}
cor(x = fHouse$Pro.Strong.Leader, y = fHouse$FHTotalAggr, use = "pairwise.complete.obs", method = "pearson")
```



```{r exercise-correlation-democracylevel-strongleader-check}
grade_code()
```

What can we conclude? There is support for our hypothesies: There is a small magnitude negative correlation, approximately -0.283. Higher levels of support for a strong leader are  associated with lower levels of institutional democracy (and vice versa).

Finally, let's examine the relationship between **wage_growth** and **dem2p_vote_share** in the data frame **counties**.  **wage_growth** is an interval-level variable that records the rate at which wages grew (or fell) over the last four years in a county. **dem2p_vote_share** is an interval-level variable measuring the share of the vote for the Democrat, Hillary Clinton, in a county in 2016. This is a case where the outcome variable is clearly identifiable: **dem2p_vote_share**. So we will place it on the y-axis. Before creating the scatter plot, filter out any missing values. Make the points "blue" and add an "lm" line to the plot without standard errors (ser `se = FALSE`). Make the line "red". In the code below I've added the labels for you.

**Explore: Plot wage growth and Democratic vote share**

```{r exercise-scatterplot-wagegrowth-democracylevel, exercise=TRUE, exercise.cap="Interactive exercise: Fill in the XXXX values to create a scatter plot of wage_growth and dem2p_vote_share, including filtering missing values and adding a fitted line. Click Submit Answer to check correctness. Three hints are available.", fig.alt="Scatter plot showing wage growth (x-axis) and Democratic vote share (y-axis) across U.S. counties in 2016, with blue points and a red fitted regression line. The points are densely clustered near low wage growth values with wide variation in vote share, while the line slopes slightly upward, indicating a weak positive relationship between wage growth and support for Clinton. The plot demonstrates how scatter plots can reveal modest trends even when data are highly dispersed."}
counties %>%
  XXXX
  XXXX
  XXXX
  XXXX
  labs(title = "Wage Growth and Support for Hillary Clinton",
       subtitle = "US Counties 2016",
       x = "Wage Growth", 
       y = "Democratic Vote Share", 
       caption = "Source: Linn, Nagler, and Zilinsky")
```

```{r exercise-scatterplot-wagegrowth-democracylevel-hint-1}
# Hint 1 of 3
You are examining the relationship between wage growth and vote share. 
Before plotting, make sure the data you pass through the pipe includes 
only rows where both variables are observed, using filter() to remove 
missing values.
```

```{r exercise-scatterplot-wagegrowth-democracylevel-hint-2}
# Hint 2 of 3
After filtering, pipe the data into ggplot().

For a scatter plot, put the explanatory variable on the x-axis 
and the outcome variable on the y-axis. Map those variables 
inside aes().
```

```{r exercise-scatterplot-wagegrowth-democracylevel-hint-3}
# Hint 3 of 3
Use + to add a point layer and a fitted line.

Set the point color inside geom_point(), and in geom_smooth() 
set the line color, se = FALSE, and method = "lm".
```


```{r exercise-scatterplot-wagegrowth-democracylevel-solution}
counties %>%
  filter(!is.na(wage_growth), !is.na(dem2p_vote_share)) %>%
ggplot(mapping = aes(x = wage_growth, y = dem2p_vote_share)) +
 geom_point(color = "blue") +
  geom_smooth(color = "red", se = FALSE, method = "lm") +
  labs(title = "Wage Growth and Support for Hillary Clinton",
       subtitle = "US Counties 2016",
       x = "Wage Growth", 
       y = "Democratic Vote Share", 
       caption = "Source: Linn, Nagler, and Zilinsky")
```

```{r exercise-scatterplot-wagegrowth-democracylevel-check}
grade_code()
```


This plot shows many counties grouped with low positive growth and that these counties had a wide range of vote share for Hillary Clinton.

Now let's estimate the correlation coefficient between these two variables.

**Practice: Calculate correlation for wage growth and vote share**

```{r exercise-correlation-wagegrowth-demvoteshare, exercise=TRUE, exercise.cap="Interactive exercise: Write code to calculate the Pearson correlation between wage_growth and dem2p_vote_share, handling missing values appropriately. Click Submit Answer to check correctness. One hint is available."}

```

```{r exercise-correlation-wagegrowth-demvoteshare-hint-1}
# Hint 1 of 1
Use cor() to measure the relationship between the two variables.

The x and y arguments are just the two variables named in the question, 
written using counties$variable. Unlike a scatter plot, it does not matter 
which variable goes in x or y.

Set method = "pearson" and choose the use option that handles missing values appropriately.
```

```{r exercise-correlation-wagegrowth-demvoteshare-solution}
cor(x = counties$wage_growth, y = counties$dem2p_vote_share, use = "pairwise.complete.obs",  method = "pearson")
```

```{r exercise-correlation-wagegrowth-demvoteshare-check}
grade_code()
```

The correlation is small and positive indicating that there is a weak, positive correlation between wage growth and the share of the vote for Hillary Clinton.

## The Takeaways

The key things to remember about examining the relationship between variables measured at the interval level are:

1. We  can summarize the relationship between two interval variables using [Pearson's correlation coefficient]{.important-text}. The correlation coefficient ranges between -1 and +1. Negative values indicate an inverse relationship between two variables: higher values of one variable are associated with lower variables of another. Positive values indicate a direct relationship: higher values of one variable are associated with higher variables of the other variable (and lower values with lower values).  The closer the correlation coefficient is to zero, the weaker the relationship. It is important to remember that the correlation coefficient assesses the strength and direction of a linear relationship. Not all relationships are linear. If a scatter plot reveals a relationship that is not linear, the correlation coefficient is meaningless.
2. We can visualize the relationship between two interval variables with a scatter plot using `ggplot()` in the [ggplot2]{.package-name} package using  `geom_point()`. We can add a line to the plot that goes through the points such that it is as close to all the points as possible using the `geom_smooth()` function.  If the line is downward sloping (goes from upper left to lower right), the correlation will be negative: higher values of one variable are associated with lower values of the other variable. If the line is upward sloping (goes from lower left to upper right),  the correlation will be positive: higher values of one variable are associated with higher values of the other variable. Finally, when the line is flat, there is no correlation between the two variables.  The steepness of the line is positively related to the strength of the correlation.



## Test your Knowledge

We've now completed three tutorials on describing relationships between two variables. Test your knowledge by answering the following questions.


```{r quiz-variable-relationship, echo=FALSE}
question("Which of the following statements is true",
  answer("We describe the relationship between two categorical variables by comparing the mean of one variable for categories of the other variable", correct=FALSE, message = "We do not summarize a caetgorical variable with a mean, so this is inapproprite. "),
  answer("A cross-tab illustrates the relationship between two categorical variables", correct=TRUE, message="Cross-tabs show the distribution of one categorical variable given values of another categorical variable. "),
  answer("The relationship between two interval variables is illustrated by comparing their distributions across categories. ", correct=FALSE, message = "We illustrate the relationship between interval variables using scatter plots. "),
  allow_retry = TRUE,
  try_again = " "
)
```



```{r quiz-correlation, echo=FALSE}
question("Which of the following statements are true",
  answer("If Pearson's correlation coefficient, $r$, is near zero, there is evidence that the variables are not related to each other. ", correct=TRUE, message="The closer a correlation is to zero, the less strongly the two variables are related. "),
  answer("If Pearson's correlation coefficient, $r$, =0.6, there is a fairly strong tendency for values of x above the mean to occur with values of y above the mean and values of x below the mean to occur with values of y above the mean. ", correct = FALSE, message="A positive correlation occurs when values of x & y above the mean tend to toccur together AND values of x & y below the mean tend to occur together . "),
  answer("If x is hours spent studying and y is correct answers to homework, we expect a negative correlation", correct = FALSE, message="Since studying should improve performance higher x should be systematically related to higher y and thus we expect a positive correlation. "),
  answer("If x is your skill at golf and y is your golf score, we expect a negative correlation", correct=TRUE, message = "The more skilled at golf, the lower your score and thus we expect a negative correlation."),
  allow_retry = TRUE,
  try_again = "Hint: Select all true statementss."
)
```

Consider the following cross-tab between the effective number of parliamentary parties in a country, **enpp3_democ**, and whether or not the country is a member of the European Union, **eu**, and answer the question:

```{r crosstable-democracy-eu, echo=FALSE}
library(descr)
CrossTable(x = world$eu, y = world$enpp3_democ, prop.r = FALSE, prop.c = TRUE, prop.t = FALSE, prop.chisq = FALSE)
```

```{r quiz-democracy-eu, echo=FALSE}
question("Which of the following statements are true",
  answer("Of countries with 1-3 political parties, almost 81% are not members of the European Union", correct = TRUE, message="There are 31 countries with 1-3 parties and 25 or 81% are not members of the EU. "),
  answer("Of EU member states, about 30% have 4-5 parties", correct=FALSE, message = "30.4 is the percent of countries with 4-5 parties who are members of the EU (69.6% are not). "),
  answer("The more parties a country has, the more likely they are to be members of the EU", correct = TRUE, message="As we move from fewer to more parties, the proportion of EU member states increases. "),
  allow_retry = TRUE,
  try_again = "Hint: More than one statement is true."
)
```

Consider the following scatter plot between government effectiveness (0 to 100 where 100 is more effective), **effectiveness**, and fertility rate (the number of children born per woman), **fertility**.

```{r scatterplot-governmenteffectiveness-fertility, echo=FALSE, warning=FALSE, message=FALSE, fig.alt="Scatter plot showing government effectiveness (x-axis) and fertility rate (y-axis) across countries, with black points and a dashed fitted regression line. The points form a clear downward trend, where higher government effectiveness is associated with fewer children per woman, indicating a strong negative relationship. The plot demonstrates how scatter plots and trend lines reveal inverse relationships between two continuous variables."}
ggplot(data = world, mapping = aes(x = effectiveness, y = fertility)) +
  geom_point() +
  geom_smooth(method = "lm", se = FALSE, linetype = "dashed") +
  labs(title = "Government Effectiveness and Fertility Around the World",
       x = "Government Effectiveness Score",
       y = "Number of Children")
```



```{r quiz-governmenteffectiveness-fertility, echo=FALSE}
question("Which of the following statements are true",
  answer("Each point represents a country's effectiveness on the x-axis and fertility on the y axis. ", correct=TRUE),
  answer("There is a positive correlation between effectiveness and fertility ", correct = FALSE, message="Because higher values of government effectiveness are associated with lower values of fertility, the correlation is negative. "),
  answer("The plot suggests a Pearson correlation close to zero", correct = FALSE, message="Randomly scattered points would suggest a correlation close to zero, but these are not randomly distributed. "),
  answer("There is a negative correlation between effectiveness and fertility. ", correct=TRUE, message = "Because higher values of government effectiveness are associated with lower values of fertility, the correlation is negative. "),
  allow_retry = TRUE,
  try_again = "Hint: More than one statement is true."
)
```

```{r tutorial-success-banner-09, results="asis", echo=FALSE}
cat('
<p role="status"
   aria-label="Superb work! You have completed this tutorial and can now describe relationships between two interval variables."
   style="font-size: 28px; font-weight: 600; color: #2e7d32; text-align: center; margin-top: 50px;">
  <span aria-hidden="true">📈 Superb work! 📈</span>
</p>
')
```
